## Práctica: Integración, normalización y detección de outliers para el algoritmo KNN

- [Pre-Procesamiento](#Pre-Procesamiento)
- [Problemas de regresión con el algoritmo KNN](#Problemas-de-regresión-con-el-algoritmo-KNN)
- [Problemas de clasificación con el algoritmo KNN](#Problemas-de-clasificación-con-el-algoritmo-KNN)

```python
# %load ../../standard_import.txt
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt

from nose.tools import assert_equal

pd.set_option('display.notebook_repr_html', False)
pd.set_option('display.max_columns', None)
pd.set_option('display.max_rows', 150)
pd.set_option('display.max_seq_items', None)
 
#%config InlineBackend.figure_formats = {'pdf',}
%matplotlib inline  
```

### Pre-Procesamiento

En esta primera parte de la práctica vamos a trabajar con un problema de regresión. En concreto vamos a trabajar con el problema de la predicción del valor de una casa en base a su superficie y a su número de habitaciones. 

Para ello tenemos dos ficheros con históricos de casas:
* El primero es ex4data1.txt y contiene casas cuya superficie está medida en pies cuadrados junto con su número de habitaciones y la valoración.
* El segundo es ex4data2.txt y contiene casas cuya superficie está medida en metros cuadrados junto con su número de habitaciones y la valoración.

Leer los datos almacenados en ambos ficheros mediante la función loadtxt de NumPy. Almacenar los datos en dos variables diferentes. 

<!-- grade_id: cell-46f046a62fe78c8a, points: 0, grade: false, solution: true -->
```python
# datos = #<RELLENAR>
# datos2 = #<RELLENAR>

```

Según la descripción del contenido de los ficheros podemos ver que existe un problema de escala ya que unas casas están medidas en pies cuadrados y otras en metros cuadrados. Un metro cuadrado equivale a 10,7639 pies cuadrados. 

Realiza una función que reciba un dataset y el índice de la columna (variable) a tratar. La función debe devolver el dataset en el que los datos de la columna correspondiente han sido transformados de pies cuadrados a metros cuadrados.

<!-- grade_id: cell-6af4c27810c647d3, points: 0, grade: false, solution: true -->
```python
def transformaSuperficie(datos, columna):
    datosEscalados = np.copy(datos)
    #<RELLENAR>

    return datosEscalados
```

Implementa una función que realice la integración de los dos conjuntos de datos anteriores y devuelva un solo conjunto de datos con los datos en la escala correcta.

La función vstack de numpy concatena arrays de forma vertical. Recibe como parámetro de entrada una tupla con los arrays a concatenar.

<!-- grade_id: cell-fc8b5b8f945e50d3, points: 0, grade: false, solution: true -->
```python
def integracion(datos, datos2):
    # transormar los datos del primer conjunto de datos
#     datosEscalados = #<RELLENAR>

    # concatenar los dos conjuntos de datos
#     datosAux = #<RELLENAR>

    return datosAux
```

También se puede observar que el rango de las dos variables es muy diferente. Por este motivo se debe realizar un proceso de normalización de los datos para que todas las variables de entrada tengan el mismo rango y, por tanto, la misma importancia.

Vamos a realizar dos formas de relizar la normalización. La primera de ellas es el método basado en la media y la desviación estándar. Se debe implementar una función que realice el siguiente proceso:
* Calcular la media de los valores de todas las variables para todas las instancias
* Calcular la desviación estándar de los valores de todas las variables para todas las instancias
* A cada valor se le debe restar la media y el resultado dividirlo por la desviación estándar
* Devolver los datos normalizados así como las medias y las desviaciones estándar calculadas

<!-- grade_id: cell-bcdee77788381452, points: 0, grade: false, solution: true -->
```python
def normalizarMediaStd(datos):
    datosEstandarizados = np.copy(datos)
#     media = #<RELLENAR>
#     desv = #<RELLENAR>
#     datosEstandarizados = #<RELLENAR>

    return datosEstandarizados, media, desv
```

Para normalizar nuevos datos de test se debe realizar el mismo proceso que en el caso anterior pero, como son datos nuevos (datos de test), tenemos que utilizar la media y la desviación estándar obtenidas con los datos de entrenamiento. De esta forma la normalización se realizará para todos los datos en las mismas condiciones. 

Definir una función que normalice nuevos datos de test. Para ello se aplica la misma ecuación que para la función anterior pero utilizando la media y la desviación estándar pasadas como parámetros de entrada (serán las obtenidas con el conjunto de entrenamiento). La función devuelve los nuevos datos normalizados.

<!-- grade_id: cell-6df3aecf4e9c9575, points: 0, grade: false, solution: true -->
```python
def normalizarMediaStdTest(datos, media, desv):
    datosEstandarizados = np.copy(datos)
#     datosEstandarizados = #<RELLENAR>

    return datosEstandarizados
```

La librería Scikit-Learn de Python nos provee una clase que realiza un proceso de normalización similar (usa la varianza en lugar de la desviación estándar). Dicha clase está dentro de la librería Preprocessing. La clase en concreto que realiza este proceso es la llamada StandardScaler, cuya información se puede ver en la siguiente URL: http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.StandardScaler.html#sklearn.preprocessing.StandardScaler. 

Para crear un objeto de esta clase es suficiente con llamar al constructor: preprocessing.StandardScaler(). Esta llamada devuelve un objeto con el que podemos llamar a varios métodos, los que vamos a utilizar son:
* fit: recibe como parámetro de entrada los datos de entrenamiento y calcula la media y la desviación estándar para cada variable. La media queda almacenada en el campo mean_ y la varianza en el campo var_.
* transform: recibe como parámetro de entrada los datos a normalizar y devuelve los datos normalizados utilizando los valores aprendidos previamente y almacenados en mean_ y var_.
* fit_transform: realiza el proceso de ambas funciones previamente descritas en una sola función.

```python
from sklearn import preprocessing

normalizarScikitPorMedia = preprocessing.StandardScaler()
```

El segundo método de normalización es el basado en el mínimo y el máximo de los datos de entrenamiento. Se debe implementar una función que realice el siguiente proceso:
* Calcular el mínimo de los valores de todas las variables para todas las instancias
* Calcular el máximo de los valores de todas las variables para todas las instancias
* A cada valor se le debe restar el mínimo y el resultado dividirlo por la diferencia entre el máximo y el mínimo
* Devolver los datos normalizados así como los mínimos y los máximos calculados

<!-- grade_id: cell-05a316dd37249695, points: 0, grade: false, solution: true -->
```python
def normalizarMaxMin(datos):
    datosEstandarizados = np.copy(datos)
#     minimos = #<RELLENAR>
#     maximos = #<RELLENAR>
#     datosEstandarizados = #<RELLENAR>

    return datosEstandarizados, minimos, maximos
```

Definir una función que normalice nuevos datos de test utilizando los mímimos y los máximos. Para ello se aplica la misma ecuación que para la función anterior pero utilizando los mínimos y los máximos pasados como parámetros de entrada (serán los obtenidos con el conjunto de entrenamiento). La función devuelve los nuevos datos normalizados.

<!-- grade_id: cell-62e7d6e0670d5cbd, points: 0, grade: false, solution: true -->
```python
def normalizarMaxMinTest(datos, minimos, maximos):
    datosEstandarizados = np.copy(datos)
#     datosEstandarizados = #<RELLENAR>

    return datosEstandarizados
```

La librería Scikit-Learn de Python también nos provee una clase que realiza este proceso de normalización. La clase en concreto que realiza este proceso es la llamada MinMaxScaler, cuya información se puede ver en la siguiente URL: http://scikit-learn.org/stable/modules/generated/sklearn.preprocessing.MinMaxScaler.html#sklearn.preprocessing.MinMaxScaler. 

Para crear un objeto de esta clase es suficiente con llamar al constructor: preprocessing.MinMaxScaler(). Esta llamada devuelve un objeto con el que podemos llamar a los mismos métodos explicados para la clase StandardScaler. En este caso el valor mínimo para cada variable (calculado con la función Fit) queda almacenado en el campo data_min_ y el máximo en el campo data_max_.

```python
normalizarScikitPorMinMax = preprocessing.MinMaxScaler()
```

El último proceso de preparación de datos que vamos a realizar en esta práctica es el de la detección de outliers. Para ello vamos a utilizar el método de la media y la desciación estándard explicado en clase de teoría. 

Definir una función que reciba los datos sobre los que detectar los outliers, los valores de las medias de todas las variables, los valores de las desviaciones estándar de todas las variables y el valor de la variable k. Esta función debe devolver una lista con los índices de aquellos ejemplos que contengan algún outlier en alguna variable. Recordar que un valor es considerado outlier si:
* valor < media - k*desviacionEstandar
* valor > media + k*desviacionEstandar

<!-- grade_id: cell-c3ec551768838512, points: 0, grade: false, solution: true -->
```python
def deteccionOutliers(datos, medias, stds, k=2.0):
    # lista con los indices de los ejemplos con outliers
    indicesOutliers = []
    # blucle para recorrer los ejemplos
    for i, ejemplo in enumerate(datos):
        # bucle para recorrer las variables
        for j in range(len(ejemplo)):

    return indicesOutliers
```

### Problemas de regresión con el algoritmo KNN

A continuación vamos a aplicar todas las funciones definidas anteriormente para realizar la fase de solucionar el problema de la valoración de las casas. En primer lugar se deben crear las matrices con los datos de entrada y de salida para los dos conjutnos de dato leídos al comienzo de la práctica.

<!-- grade_id: cell-8c8580776a1d4674, points: 0, grade: false, solution: true -->
```python
# datos de entrada y salida del primer fichero
# X1 = #<RELLENAR>
# y1 = #<RELLENAR>

# datos de entrada y salida del segundo fichero
# X2 = #<RELLENAR>
# y2 = #<RELLENAR>

```

Ahora debéis integrar los dos conjuntos de datos para formar un solo conjunto de datos con el que aplicaremos el resto de procesos.

<!-- grade_id: cell-00534356d8993875, points: 0, grade: false, solution: true -->
```python
# datos de entrada concatenados e integrados
# X = #<RELLENAR>

# Calcular el número de filas del conjunto de datos integrado
numFilas = X.shape[0]
# concatenar las salidas de los dos conjutnos de datos
y = np.vstack((y1.reshape(-1,1),y2.reshape(-1,1)))
```

<!-- grade_id: cell-1712beaedfc24505, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(numFilas, 57)
```

El siguiente paso es obtener los índices de aquellos ejemplos que contengan outliers. Para ello debéis llamar a la función que normaliza los datos con el método de la media y la desviación estándar para obtener las medias y las desviaciones estándar de cada variable. Después debéis llamar a la función que obtiene los índices de los ejemplos con outliers y almacenarlos en una variable.

<!-- grade_id: cell-6ce9f77fcef82b63, points: 0, grade: false, solution: true -->
```python
# Cálculo de las medias y desviaciones
# Xnorm, medias, stds = #<RELLENAR>

# Llamada a la función que obtiene los índices de los ejemlos con outliers
# indicesOutliers = #<RELLENAR>

print(indicesOutliers)
```

<!-- grade_id: cell-9dc20679e5e65606, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: ind, list(indicesOutliers))), [13, 55, 56])
```

El siguiente paso es tratar los ejemplos que contengan outliers. En este caso, se ha decidido eliminarlos del conjunto de datos. Por este motivo, debéis crear un nuevo conjunto de datos en el que se hayan eliminado los outliers.

La función delete de numpy borra filas de un array (np.delete(array, indices, axis=0))

<!-- grade_id: cell-52e8d17bda18c047, points: 0, grade: false, solution: true -->
```python
# Borrado de los ejemplos con outliers
# XsinOutliers = #<RELLENAR>
# YsinOutliers = #<RELLENAR>

numFilas = XsinOutliers.shape[0]
```

<!-- grade_id: cell-aede5ab85c2f8e7e, points: 1, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(numFilas, 54)
```

En este momento ya tenemos los datos integrados y sin outliers por lo que podemos aplicar la técnica de aprendizaje automático. En este caso vamos a aplicar el algoritmo de los K vecinos más cercanos para problemas de regresión. La librería Scikit-Learn nos ofrece una implementación de dicho método que está dentro del paquete neighbors y cuya clase específica es KNeighborsRegressor, cuya información se puede consultar en la sigueinte URL: http://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsRegressor.html#sklearn.neighbors.KNeighborsRegressor

Lo primero que se debe realizar es llamar a la función que crea el modelo de regresión. Esta función tiene varios parámetros que determinan el comprtamiento del algoritmo. La llamada al constructor y sus parámetros son los siguientes:

modeloRegresion = neighbors.KNeighborsRegressor(n_neighbors  = K, weights = tipoVoto, metric = tipoDistancia, p = r)

Los diferentes parámetros de entrada son
* n_neighbors = K: número de vecinos a considerar (valor por defecto = 5)
* weights = tipoVoto: forma de votar (peso de cada ejemplo cercano). tipoVoto puede tomar los siguientes valores:
    * 'uniform': voto por mayoría  (valor por defecto)
    * 'distance': voto en función de la inversa de la distancia
* metric = tipoDistancia: forma de calcular la distancia entre los ejemplos.  tipoDistancia puede tomar los siguientes valores:
    * 'manhattan': distancia de manhattan
    * 'euclidean': distancia euclidea
    * 'minkowski': distancia de Minkowski (valor por defecto)
* r: en caso de utilizar la distancia de Minkowski hay que especificar el valor del parámetro p que se corresponde al exponente r visto en la clase de teoría. r puede cualquier valor, entre ellos:
    * r = 1: distancia de manhattan
    * r = 2: distancia euclidea (valor por defecto)
    
Crear el modelo inicial con todos los valores por defecto. Almacenar el modelo en una variable.


```python
from sklearn import neighbors

modeloReg = neighbors.KNeighborsRegressor()
```

Una vez creado el modelo debemos entrenarlo. Para ello se debe llamar al método fit del objeto creado anteriormente. A dicho método se le deben pasar los datos de entrenamiento (en este caso todos los datos disponibles) tanto de las variables de entrada (sin outliers) como de la salida.

<!-- grade_id: cell-c9fba90bd3a43a95, points: 0, grade: false, solution: true -->
```python
# modeloReg = modeloReg.fit(#<RELLENAR>)

```

Una vez que el modelo está entrenado lo podemos utilizar para realizar predicciones de nuevos datos. Para ello se debe llamar al método predict al que se pasa comp parámetro de entrada los ejemplos a predecir (en este caso vamos a predecir el valor de todas las casas disponibles sin outliers). Guardar el resultado de las predicciones en una variable.

<!-- grade_id: cell-af615ce3963c1dbf, points: 0, grade: false, solution: true -->
```python
# predicciones = modeloReg.predict(#<RELLENAR>) 

```

En este punto ya tenemos las predicciones de todas las casas por lo que podemos calcular el rendimiento del modelo. Para ello debéis programar la ecuación que calcula el error cuadrático medio de las predicciones realizadas anteriormente.

<!-- grade_id: cell-7947687134b89396, points: 0, grade: false, solution: true -->
```python
# error = #<RELLENAR>

print(error)
```

<!-- grade_id: cell-bec6b5e86b2c5cd8, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(error, 2), 3534488808.38)
```

La librería Scikit-Learn ofrece una función que calcula el error cuadrático medio de las predicciones sin tener que calcularlo. Esta clase está dentro del paquete metrics y se llama mean_squared_error. A esta función se le pasa como argumentos de entrada los valores reales de los ejemplos y las predicciones realizadas por el modelo. Obviamente, devuelve el error cuadrático medio de las predicciones realizadas.

Utilizar dicha función para calcular el error cuadrático medio de las predicciones realizadas anteriormente.

<!-- grade_id: cell-250f1f2ba4dab1d4, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error

# errorScikit = mean_squared_error(#<RELLENAR>)

print(errorScikit)
```

<!-- grade_id: cell-403ab1c84bb1ddc8, points: 0, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(errorScikit, 2), 3534488808.38)
```

Una vez que sabemos utiizar la librería de Scikit-Learn para aplicar el algoritmo KNN de regresión vamos a estudiar la influencia de la normalización de los datos a la hora de realizar las predicciones. En primer lugar vamos a realizar la normalización por el método de la media y la desviación estándar.

Debéis realizar lo siguiente:
* Llamar a la función que hemos implementado para normalizar los datos (sin outliers) según este método.
* Realizar el aprendizaje con los ejemplos normalizados (K=5)
* Realizar la predicción de los ejemplos normalizados
* Calcular el error cuadrático medio

¿Es mayor o menor que el obtenido sin normalizar los datos?

<!-- grade_id: cell-f3f5c30fb2fc1732, points: 0, grade: false, solution: true -->
```python
# Xnorm, media, desv = #<RELLENAR>
modeloReg = neighbors.KNeighborsRegressor()
# modeloReg = modeloReg.fit(#<RELLENAR>)
# predicciones = modeloReg.predict(#<RELLENAR>) 
# errorScikit = mean_squared_error(#<RELLENAR>)

print(errorScikit)
```

<!-- grade_id: cell-4c4653fc2b4c6bcc, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(errorScikit, 2), 3473398814.02)
```

Realizar el mismo proceso que antes pero utilizando la normalización por el máximo y el mínimo.

<!-- grade_id: cell-cdd96e3bd66c21fd, points: 0, grade: false, solution: true -->
```python
# Xnorm, minimo, maximo = normalizarMaxMin(#<RELLENAR>)
modeloReg = neighbors.KNeighborsRegressor()
# modeloReg = modeloReg.fit(#<RELLENAR>)
# predicciones = modeloReg.predict(#<RELLENAR>) 
# errorScikit = mean_squared_error(#<RELLENAR>)

print(errorScikit)
```

<!-- grade_id: cell-fd091717ad21e818, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(errorScikit, 2), 3530600295.50)
```

### Problemas de clasificación con el algoritmo KNN

En esta segunda parte de la práctica vamos a trabajar con un problema de clasificación. En concreto vamos a trabajar con el problema de la predicción de si un microchip pasa los tests de calidad o no

Leer los datos almacenados en el fichero ex4data3.txt mediante la función loadtxt de NumPy. Almacenar los datos en una variable.

<!-- grade_id: cell-4c6ef51d2678006c, points: 0, grade: false, solution: true -->
```python
# datos3 = #<RELLENAR>

```

Crear la matriz con los datos de entrada (X) y de salida (y).

<!-- grade_id: cell-d91d1b9da096c7b4, points: 0, grade: false, solution: true -->
```python
# X = #<RELLENAR>
# y = #<RELLENAR>

```

Normalizar los datos de entrada con el método de la media y la desviación estándar ofrecito por la librería Scikit-Learn. Mostrar los valores de los 5 primeros ejemplos.

<!-- grade_id: cell-4d01fcd890d8847e, points: 0, grade: false, solution: true -->
```python
normalizarScikitPorMedia = preprocessing.StandardScaler()
# Xnorm = normalizarScikitPorMedia.fit_transform(#<RELLENAR>)

print(Xnorm[:5,:])
```

En este caso, al ser un problema de clasificación, debemos aplicar el algotimo KNN para problemas de clasificación. La librería Scikit-Learn nos ofrece una implementación de dicho método que está dentro del paquete neighbors y cuya clase específica es KNeighborsClassifier, cuya información se puede consultar en la sigueinte URL: http://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html#sklearn.neighbors.KNeighborsClassifier

Esta clase funciona exactamente igual que la que soluciona problemas de regresión.
    
Crear el modelo inicial con todos los valores por defecto. Almacenar el modelo en una variable.

```python
modeloClas = neighbors.KNeighborsClassifier()
```

El entrenamiento (fit) y la predicción de nuevos ejemplos (predict) se realizan igual que antes.

Sin embargo, para obtener el rendimiento del clasificador debemos aplicar otras funciones de la librería metrics. En concreto vamos a aplicar:
* accuracy_score: los parámetros de entrada son las clases reales y las predicciones y devuelve el ratio de ejemplos correctamente clasificados.
    * URL: http://scikit-learn.org/stable/modules/generated/sklearn.metrics.accuracy_score.html#sklearn.metrics.accuracy_score
* confusion_matrix: los parámetros de entrada son las clases reales y las predicciones y devuelve la matriz de confusión.
    * URL: http://scikit-learn.org/stable/modules/generated/sklearn.metrics.confusion_matrix.html#sklearn.metrics.confusion_matrix
* classification_report: los parámetros de entrada son las clases reales y las predicciones y devuelve varias medidas de rendimiento como la precisión, el recall, la medida F1 y el soporte (número real de ejemplos de cada clase).
    * URL: http://scikit-learn.org/stable/modules/generated/sklearn.metrics.classification_report.html#sklearn.metrics.classification_report

Entrenar el clasificador con los datos normalizados, realizar sus predicciones y mostrar el resultado de las tres funciones de rendimiento anteriores.


<!-- grade_id: cell-e3fbd03b019d00d1, points: 0, grade: false, solution: true -->
```python
from sklearn import metrics

# Entrenamineto del modelo
# modeloClas = #<RELLENAR>
# Predicciones de los datos
# predicciones = #<RELLENAR>
# Cálculo del accuracy
# accuracy = metrics.accuracy_score(#<RELLENAR>)

print(accuracy)

# Cálculo de la matriz de confusión
# matrizConfusion = metrics.confusion_matrix(#<RELLENAR>)

print(matrizConfusion)

# print metrics.classification_report(#<RELLENAR>)

```

<!-- grade_id: cell-e1a9e5e5dd024985, points: 4, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(accuracy, 5), 0.84746)
assert_equal(list(map(lambda ind: ind, list(matrizConfusion.ravel()))), [48, 12, 6, 52])
```

Ahora, vamos analizar el efecto de los parámetros del algoritmo KNN en los resultados obtenidos (accuracy rate):
* Número de vecinos (parámetro n_neighbors)
    * Probar los resultados con 1, 3, 5 y 7 vecinos

<!-- grade_id: cell-be2221d52a2ac2eb, points: 0, grade: false, solution: true -->
```python
# Lista con los accuracies por cada valor k
accScikit = []
for k in [1,3,5,7]:
    # Creación del modelo
#     modeloClas = neighbors.KNeighborsClassifier(#<RELLENAR>)
    # Entrenamiento del modelo
#     modeloClas = #<RELLENAR>
    # Predicciones de los datos
#     predicciones = #<RELLENAR>
    # Cálculo del accuracy rate
#     accuracy = #<RELLENAR>
    # Se añade el accuracy a la lista
#     accScikit.append(accuracy)

    # Se muestra el número de vecinos y el accuracy
    print(k, accuracy)
```

<!-- grade_id: cell-8a0ae49f8ba35b9a, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: round(ind, 3), list(accScikit))), [1.0, 0.856, 0.847, 0.831])
```

Utilizando un número de vecinos igual a 3, vamos a analizar el efecto de la forma de votación.
* Tipo de voto (parámetro weights)
    * Probar los valores 'uniform', 'distance'

<!-- grade_id: cell-862998213fac8383, points: 0, grade: false, solution: true -->
```python
# Lista con los accuracies por cada tipo de voto
accScikit = []
for tipoVoto in ['uniform', 'distance']:
    # Creación del modelo
#     modeloClas = neighbors.KNeighborsClassifier(#<RELLENAR>)
    # Entrenamiento del modelo
#     modeloClas = #<RELLENAR>
    # Predicciones de los datos
#     predicciones = #<RELLENAR>
    # Cálculo del accuracy rate
#     accuracy = #<RELLENAR>
    # Se añade el accuracy a la lista
#     accScikit.append(accuracy)

    # Se muestra el tipo de voto y el accuracy
    print(tipoVoto, accuracy)
```

<!-- grade_id: cell-bc53b64c713179d7, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: round(ind, 3), list(accScikit))), [0.856, 1.0])
```

Utilizando un número de vecinos igual a 3, vamos a analizar el efecto de la forma de votación.
* Tipo de distancia (parámetro metric)
    * Probar los valores 'manhattan', 'euclidean'

<!-- grade_id: cell-1f16fbee2fd5e4bf, points: 0, grade: false, solution: true -->
```python
# Lista con los accuracies por cada tipo de voto
accScikit = []
for tipoDistancia in ['manhattan', 'euclidean']:
    # Creación del modelo
#     modeloClas = neighbors.KNeighborsClassifier(#<RELLENAR>)
    # Entrenamiento del modelo
#     modeloClas = #<RELLENAR>
    # Predicciones de los datos
#     predicciones = #<RELLENAR>
    # Cálculo del accuracy rate
#     accuracy = #<RELLENAR>
    # Se añade el accuracy a la lista
#     accScikit.append(accuracy)

    # Se muestra el tipo de distancia y el accuracy
    print(tipoDistancia, accuracy)
```

<!-- grade_id: cell-bdb6a773a5b5256f, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: round(ind, 3), list(accScikit))), [0.881, 0.856])
```

Por último, vamos a visualizar la frontera de decisión que genera el clasificador KNN. Debéis hacer lo siguiente:
* Aprender el clasificador con los datos normalizados utilizando 3 vecinos, y la distancia de Manhattan. (A realizar)
* Se crea una nube de puntos con todas las combinaciones entre el mínimo (-0.1) y el máximo (+0.1) con incrementos de 0.02 de las dos variables.
* Se realiza la predicción de todos los ejemplos generados.
* Se crea la gráfica donde la predicción de cada clase sale en dieferentes colores (contour) y se muestran los ejemplos de entrenamiento (scatter)

<!-- grade_id: cell-ce23f6ee08a044f7, points: 0, grade: false, solution: true -->
```python
# Creación del modelo
# modeloClas = #<RELLENAR>
# Entrenamiento del modelo
# modeloClas = #<RELLENAR>


#se crea una grñafca para mostrar la superficie de decisión del clasificador aprendido
h = .02  # tamaño de avance en el mesh
#TAREA: calcula el mínimo (menos 1) y el máximo (más 1) de las variables x e y
x_min = Xnorm[:, 0].min() - 0.1
x_max = Xnorm[:, 0].max() + 0.1
y_min = Xnorm[:, 1].min() - 0.1
y_max = Xnorm[:, 1].max() + 0.1

# Se crean todas las combinaciones de valores
xx, yy = np.meshgrid(np.arange(x_min, x_max, h), np.arange(y_min, y_max, h))
# Se predicen todos los puntos de la superficio de la gráfica con el clasificador aprendido
Z = modeloClas.predict(np.c_[xx.ravel(), yy.ravel()])

#Se muestran los resultados en la gráfica: la función contourf colorea toda la superficie
Z = Z.reshape(xx.shape)
plt.figure(1, figsize=(8,8))
plt.contourf(xx, yy, Z, cmap=plt.cm.Paired, alpha=0.5)

#Se muestran los datos del problema: la funcion scatter muestra puntos aislados y los colorea en función de su valor
plt.scatter(Xnorm[:, 0], Xnorm[:, 1], c=y, cmap=plt.cm.Paired, edgecolors='k')

#se muestra la figura
plt.show()
```