## Práctica 3 - Regresión logística

- [Regresión logística](#Regresión-logística)
- [Regresión logística con regularización](#Regresión-logística-con-regularización)

Importamos todas las librerías que vamos a utilizar durante la segunda práctica.

```python
# %load ../../standard_import.txt
import pandas as pd
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
from nose.tools import assert_equal

from scipy.optimize import minimize
from sklearn.preprocessing import PolynomialFeatures

pd.set_option('display.notebook_repr_html', False)
pd.set_option('display.max_columns', None)
pd.set_option('display.max_rows', 150)
pd.set_option('display.max_seq_items', None)
 
#%config InlineBackend.figure_formats = {'pdf',}
%matplotlib inline
```

Definimos una función que lee un archivo que contiene los datos de entrenamiento. La función recibe como parámetro de entrada el nombre del fichero que contiene los datos de entrenamiento así como el símbolo utilizado para separar los valores de las variables del problema. La función muestra las dimensiones del dataset así como los 5 primeros ejemplos del mismo.

Debéis llamar a la función loadtxt de numpy con los parámetros apropiados.

<!-- grade_id: cell-5236a224e35865d8, points: 0, grade: false, solution: true -->
```python
def lecturaDatos(nombreFichero, simboloDelimitador):
#     datos = #<RELLENAR>


    print('Dimensiones de los datos: ',datos.shape)
    print(datos[1:6,:])
    return(datos)
```

Definimos una función que muestre los datos de entrenamiento gráficamente. 

Esta función va a ser genérica para problemas de clasificación de dos clases por lo que
* La clase negativa (no admitido) va a ser la que tenga asociada la etiqueta 0
* La clase positiva (admitido) va a ser la que tenga asociada la etiqueta 1

Además asumimos que la variable de salida (clase) será la última.

La función recibe como parámetros de entrada los datos, las etiquetas de los ejes x e y (dos variables de entrada) así como la etiqueta para los datos la clase positiva y de la clase negativa (se muestran en la leyenda de la figura).

<!-- grade_id: cell-9fa20246bd923826, points: 0, grade: false, solution: true -->
```python
def muestraGrafica(datos, etiqueta_ejeX, etiqueta_ejeY, etiqueta_Positiva, etiqueta_Negativa, axes=None):
    # La variable claseNegativa debe ser un array de booleanos compuesto por tantos elementos como ejemplos
        # Aquellos alumnos no admitidos (clase 0) deben tener True en su correspondiente posicion del array de booleanos y False en el resto
    claseNegativa = datos[:,-1] == 0
    # La variable clasePositiva debe ser un array de booleanos compuesto por tantos elementos como ejemplos
        # Aquellos alumnos admitidos (clase 1) deben tener True en su correspondiente posicion del array de booleanos y False en el resto
#     clasePositiva = #<RELLENAR>


    
    # Si no se pasa un manejador de una figura se crea una
    if axes == None:
        axes = plt.gca()
    # Se muestra la nube de puntos: + negros para la clase 1 (admitidos) y circulos amarillos para la clase 0 (no admitidos)
#     axes.scatter(#<RELLENAR>, #<RELLENAR>, marker='+', c='k', s=60, linewidth=2, label=etiqueta_Positiva)
#     axes.scatter(#<RELLENAR>, #<RELLENAR>, c='y', s=60, label=etiqueta_Negativa)


    axes.set_xlabel(etiqueta_ejeX)
    axes.set_ylabel(etiqueta_ejeY)
    axes.legend(frameon= True, fancybox = True);
```

### Regresión logística

Se leen los datos utilizando la función definida anteriormente. El nombre del fichero a leer es ex3data1.txt y el delimitador de los valores de las variables es la coma (,).

Los datos contenidos en este fichero corresponden a un problema en el que se debe predecir si un estudiante es admitido en una universidad o no. El acceso a la universidad viene determinado por la nota obtenida por los estudiantes en dos exámenes realizados previamente. El archivo contiene el histórico de las notas obtenidas en ambos exámenes (variables de entrada) por 100 alumnos así como si han ingresado a la universidad o no (variable de salida).

<!-- grade_id: cell-22f6fa8af6791ebe, points: 0, grade: false, solution: true -->
```python
# datos = #<RELLENAR>


```

A continuación se crean los datos de entrada (X) y de salida (y) para poder realizar el aprendizaje de los parámetros de la regresión logística ($\theta$). Recordar que la primera columna de los datos de entrada debe estar compuesta por unos para multiplicar al término independiente de la regresión logística. Por tanto, la variable X estará compuesta por tantas columnas como variables de entrada más una (la de unos). En este caso serán 3 columnas. La variable y debe contier la información de si los alumnos han sido admitidos o no.

<!-- grade_id: cell-0ed9a3ad8c08868d, points: 0, grade: false, solution: true -->
```python
# X = #<RELLENAR>
# y = #<RELLENAR>


```

Realizar la llamada a la función que muestra los ejemplos utilizando los datos recién leídos y las etiquetas apropiadas. Los alumnos admitidos se codifican con 1 y los no admitidos con 0.

<!-- grade_id: cell-8f226a41e9f4a6b6, points: 0, grade: false, solution: true -->
```python


```

Vamos a recordar el modelo de la regresión logística. Es el mismo que para la regresión lineal pero aplicando una función sigmoide sobre el resultado para conseguir una función convexa y, de esta forma, poder aplicar el descenso por gradiente para minimizar la función de coste.

#### Hipótesis de la regresión logística
#### $$ h_{\theta}(x) = g(\theta^{T}x)$$
#### La función $g$ es la función sigmoide definida como
#### $$ g(z)=\frac{1}{1+e^{−z}} $$

Definir la función que aplica la función sigmoide, $g(z)$, sobre un dato $z$. Este dato puede ser un array de valores.

La función exp de numpy (pasando un valor como parámetro de entrada) sirve para calular $e^{valor}$

<!-- grade_id: cell-a04357a8872ab2ed, points: 0, grade: false, solution: true -->
```python
def sigmoide(z):
#     return(#<RELLENAR>)


```

Comprobar si la función sigmoide funciona correctamente, recordar que:
* Para el valor 0.0 el resultado debe ser 0.5
* Para valores grandes el resultado debe ser muy próximo a 1.0
* Para valores pequeños el resultado debe ser muy próximo a 0.0

Ejericicio:
* Comprobar el resultado para el valor 0.0
* Comprobar el resultado para un array de numpy compuesto por los valores -5.0 y 5.0

<!-- grade_id: cell-edc57fb7a385b5ce, points: 0, grade: false, solution: true -->
```python
# resultado1 = #<RELLENAR>


print(resultado1)
# resultado2 = #<RELLENAR>


print(resultado2)
```

<!-- grade_id: cell-c4b7abd08f24d17d, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(resultado1, 0.5)
assert_equal(list(map(lambda ind: round(ind, 5), list(resultado2.ravel()))), [0.00669,  0.99331])
```

La librería Scipy contiene una función que realiza exáctamente este mismo cálculo. Se puede consultar su información la siguiente URL:<BR>
http://docs.scipy.org/doc/scipy/reference/generated/scipy.special.expit.html#scipy.special.expit

Se podría utilizar en lugar de la función sigmoide que acabamos de definir.

A continuación vamos a implementar la función de coste. Vamos a recordar su expresión normal y vectorizada.

#### Función de coste
#### $$ J(\theta) = \frac{1}{m}\sum_{i=1}^{m}\big[-y^{(i)}\, log\,( h_\theta\,(x^{(i)}))-(1-y^{(i)})\,log\,(1-h_\theta(x^{(i)}))\big]$$
#### Función de coste vectorizada
#### $$ J(\theta) = -\frac{1}{m}\big((\,log\,(g(X\theta))^Ty+(\,log\,(1-g(X\theta))^T(1-y)\big)$$

Definir la función de coste vectorizada mostrada anteriormente $J(\theta)$. Esta función recibe como parámetros de entrada los valores de los parámetros $\theta$ sobre los que calcular el coste con respecto a los datos de entrada (X) y sus correspondientes salidas (y).

```python
def funcionCoste(theta, X, y):
    m = y.size
    h = sigmoide(X.dot(theta))
    
    J = -1.0*(1.0/m)*(np.log(h).T.dot(y)+np.log(1.0-h).T.dot(1.0-y))
               
    if np.isnan(J[0]):
        return(np.inf)
    return(J[0])
```

Recordamos como son las derivadas parciales para poder definir la función que aplica el descenso por gradiente y de esta forma aprender los valores de los parámetros de la regresión logística ($\theta$)

#### Derivada parcial

#### $$ \frac{\partial J(\theta)}{\partial\theta_{j}} = \frac{1}{m}\sum_{i=1}^{m} ( h_\theta (x^{(i)})-y^{(i)})x^{(i)}_{j} $$ 
#### Derivada parcial vectorizada
#### $$ \frac{\partial J(\theta)}{\partial\theta_{j}} = \frac{1}{m} X^T(g(X\theta)-y)$$


Definir la función que realiza la derivada parcial vectorizada mostrada previamente. Esta función recibe como parámetros de entrada los valores iniciales de los parámetros de la regresión logística ($\theta$) así como los valores de las variables de entrada del problema (X) y sus correspondientes salidas (y).

En la variable h tenemos el resultado de $g(X\theta)$

Ejercicio: 
* Implementa la ecuación de la derivada parcial $\frac{1}{m} X^T(g(X\theta)-y)$

<!-- grade_id: cell-c602192ee8ff954f, points: 0, grade: false, solution: true -->
```python
def derivadaParcial(theta, X, y):
    m = y.size
    h = sigmoide(X.dot(theta.reshape(-1,1)))
    
#     derParcial = #<RELLENAR>



    return(derParcial.flatten())
```

Ejercicio:
* Crear los valores iniciales de los parámetros de la regresión logística como un vector columna de ceros. 
* Comprobar el coste del modelo con los datos leídos utilizando estos valores iniciales de los parámetros $\theta$.
* Obtener los valores de la derivada parcial de los parámetros de la regresión logística llamando a la función que la realiza utilizando los datos de entrenamiento leídos anteriormente.

<!-- grade_id: cell-26477f23b9b5051e, points: 0, grade: false, solution: true -->
```python
# Creación de la columna de ceros
# thetaInicial = #<RELLENAR>


# Llamada a la funcion de calculo del coste
# costeInicial = #<RELLENAR>


# Llamada a la funcion que realiza la derivada parcial
# derParcial = #<RELLENAR>

print('Coste del modelo inicial: \n', costeInicial)
print('Derivada parcial: \n', derParcial)
```

<!-- grade_id: cell-3a7fb6b5c0c804c8, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(costeInicial, 5), 0.69315)
assert_equal(list(map(lambda ind: round(ind, 5), list(derParcial.ravel()))), [-0.1,  -12.00922, -11.26284])
```

#### Optimización de la función de coste

Para realizar el aprendizaje de los parámetros de la regresión logística vamos a utilizar la función minimize importada al comienzo de la práctica. Esta función realiza el proceso de descenso por gradiente y recibe como parámetros de entrada:
* La función que calcula el coste del modelo
* Los valores iniciales de los parámetros del modelo (variable thetaInicial)
* Los valores de los ejemplos de entrenamiento (tanto de las variables de entrada (X) como de salida (y)) en forma de tupla (parámetro args)
* La función que realiza las derivadas parciales (parámetro jac)
* El número máximo de iteraciones del descenso por gradiente (establecido a 400)

Los parámetros del modelo aprendidos ($\theta$) se devuelven en el campo x de la variable utilizada para recoger la salida de la función minimize. En este caso se pueden consultar imprimiedo res.x. 

```python
res = minimize(funcionCoste, thetaInicial, args=(X,y), jac=derivadaParcial, options={'maxiter':400})
print(res.x)
```

<!-- grade_id: cell-4d84bebdfd144e18, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: round(ind, 5), list(res.x.ravel()))), [-25.16132,  0.20623, 0.20147])
```

#### Predicción de nuevos datos

Para predecir la clase de nuevos estudiantes se debe definir una función que recibe los valores de los parámetros del modelo ($\theta$), una lista de lista con los valores de las variables de entrada de todos los alumnos (X) y un umbral para decidir si se predice la clase admitido o no admitido. La función sigmoide definida anteriormente devuelve un valor en el intervalo [0, 1] que debemos transformar a 0 o 1. Para realizar este proceso utilizamos el umbral de tal forma que si el valor devuelto por la función sigmoide es menor que el umbral se obtiene False (0) y en otro caso True (1). Por tanto, este umbral por defecto debe ser 0.5 para que ambas clases tengan la misma probabilidad de ser predichas.

Ejercicio:
* Calcular la probabilidad de admisión de todos los alumnos (función sigmoide pasando el resultado de la ecuación del modelo $X\theta^{T}$). 
    * Se asigna a la variable p.
* Comparar cada valor devuelto por la función sigomide (almacenado en p) con el umbral para obtener si son admitidos (mayor o igual que el umbral: True) o no (menor que el umbral: False). 
    * La variable p se actualiza con este resultado.

<!-- grade_id: cell-10c82245e36cda8a, points: 0, grade: false, solution: true -->
```python
def prediccion(theta, X, umbral=0.5):
#     p = #<RELLENAR>

    # La función astype la utilizamos para convertir booleanos (False (0), True (1)) en números enteros {0, 1}.
    return(p.astype('int'))
```

Calcular la probabilidad de admisión de un estudiante que ha obtenido un 45 en el primer examen y un 85 en el segundo. Para ello se debe llamar a la función sigmoide ya que nos devuelve la probabilidad de que un estudiante sea admitido (número entre 0 y 1 como se ha mencionado en el apartado previo).

Ejercicio:
* Definir el array de numpy con los datos del alumno para que pueda ser predicho por el modelo (acordaros de incluir el 1).

Nota:
* Fijaros que a la sigmoide se le pasa el alumno (X) multiplicado por los parámetros del modelo transpuestos (res.x.T).

<!-- grade_id: cell-55b696c5195f0255, points: 0, grade: false, solution: true -->
```python
# Predecir utilizando los theta obtenidos tras aplicar la función minimize (res.x)
# alumno = #<RELLENAR>

probabilidad = sigmoide(alumno.dot(res.x.T))
print(probabilidad)
```

<!-- grade_id: cell-fedf000a263c4744, points: 1, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(probabilidad, 5), 0.77629)
```

Ejercicio:
* Realizar la predicción de todos los alumnos almacenados en el fichero que contiene el histórico (el leído al comienzo de la práctica). 
    * Llamar a la función predicción definida anteriormente pasando los valores de los parámetros del modelo (res.x) y los datos a predecir.
* Obtener el porcentaje de acierto aplicando los siguientes pasos:
    * Obtener un array de booleanos comprobando si los alumnos se predicen correctamente o no (p==y.ravel())
    * Obtener el número de alumnos correctamente predichos, sumar el array anterior (con la función sum se suman los True ya que son unos)
    * Obtener el porcentaje de alumnos correctamente predichos divididiendo la suma anterior entre el número total de alumnos (p.size) y multiplicado por 100. 
        * Este resultado se almacena en la variable accuracy.

<!-- grade_id: cell-a6aa483bab2ca6fe, points: 0, grade: false, solution: true -->
```python
# p = #<RELLENAR>
# accuracy = #<RELLENAR>

print('Precisión en entrenamiento {}%'.format(accuracy))
```

<!-- grade_id: cell-d8c8b0b6ba86795f, points: 2, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(accuracy, 1), 89.0)
```

#### Frontera de decisión

Ahora vamos a mostrar la frontera de decisión que genera el modelo de regresión logística que acabamos de aprender. Recordar que la frontera de decisión se genera en el punto en el que las probabilidades del ejemplo a todas las clases son iguales y, por tanto, no se sabe determinar la clase. A partir de ese punto la clase cambia entre las que estén afectadas por la frontera (generalmente dos).

NOTA:
* Fijaros en como se obtiene la frontera de decisión (línea azul): se aplica la sigmoide (probabilidad de cada alumno) y se muestran solamente los alumnos cuya probabilidad sea 0.5.

```python
# Mostramos el punto que hemos clasificado anteriormente como un triángulo rojo
plt.scatter(45, 85, s=60, c='r', marker='v', label='(45, 85)')
muestraGrafica(datos, 'Nota en el examen 1', 'Nota en el examen 2', 'Admitido', 'No admitido')
# Calculamos el mínimo y máximo de cada variable de entrada
x1_min, x1_max = X[:,1].min(), X[:,1].max(),
x2_min, x2_max = X[:,2].min(), X[:,2].max(),
# Obtenemos todas las combinaciones de valores de las dos variables (comprendidos entre el mínimo y el máximo y espaciados linealmente)
xx1, xx2 = np.meshgrid(np.linspace(x1_min, x1_max), np.linspace(x2_min, x2_max))
# Aplicamos la función sigmoide para obtener la probabilidad de admisión de todas las combinaciones anteriores
h = sigmoide(np.c_[np.ones((xx1.ravel().shape[0],1)), xx1.ravel(), xx2.ravel()].dot(res.x))
h = h.reshape(xx1.shape)
# Mostramos con una línea azul todos las combinaciones cuya probabilidad de admisión sea 0.5 (equiprobable: no sabemos elegir)
plt.contour(xx1, xx2, h, [0.5], linewidths=1, colors='b');
```

### Regresión logística con regularización

En la segunda parte de la práctica vamos a trabajar con un nuevo problema de clasificación. En este caso se trata de predecir si un microchip pasa los tests de calidad o no. Para ello, a cada microchip se le realizan dos tests, cuyos resultados se almacenan en un archivo que contiene los datos históricos de 118 microchips. Por tanto, para cada microchip se tienen dos variables de entrada (los resultados de los dos tests) y la variable de salida que tiene dos valores: pasa el test de calidad, son aceptados, (1) o no, son rechazados, (0). Los datos de este problema están almacenados en el fichero llamado ex3data2.txt cuyo delimitador también es la coma (,).

Realiza la llamada a la función que realiza la lectura de los ficheros y almacena los datos.

<!-- grade_id: cell-0a0c26184f0530fc, points: 0, grade: false, solution: true -->
```python
# data2 = #<RELLENAR>

```

Crea las variables de entrada (X) y de salida (y).

<!-- grade_id: cell-3640287e9fced594, points: 0, grade: false, solution: true -->
```python
# X = #<RELLENAR>
# y = #<RELLENAR>

print(X.shape)
```

Llama a la función que muestra la gráfica con los ejemplos del histórico de microchips y nombra correctamente los ejes y lo títulos de la leyenda. La clase positiva (1) son los microchips aceptados y la clase negativa (0) los rechazados.

<!-- grade_id: cell-2f2e52c6becfbd24, points: 0, grade: false, solution: true -->
```python

```

#### Polinomios

Como podemos ver en la gráfica anterior los ejemplos no son linealmente separables. Por este motivo, en primer lugar, vamos a crear nuevas variables del modelo que nos permitan obtener fronteras de decisión no lineales. Para ello
* Aplicar la función PolynomialFeatures pasando como parámetro de entrada el valor 6 ya que vamos a generar todos los términos polinómicos de x1 y x2 hasta la potencia 6: en total hay 28 combinaciones (por tanto se generan 28 variables).
* Una vez creado el objeto que nos permite realizar dicha transformación llamamos a la función fit_transform utilizando como parámetro de entrada las variables de entrada de nuestro problema.
* Esta función tiene dos partes (fit y transform)
    * La parte fit aprende los polinomios para generar las nuevas variables
    * La parte transform aplica los polinomios generados para transformar los datos de entrada y que como consecuencia pasemos de tener dos variables de entrada a 28.

NOTA: los términos polinómicos de las variables a y b hasta la potencia p=3 serían: 1, a, b, $a^2$, $b^2$, $a^3$, $b^3$, ab, a$b^2$, $a^2$$b^2$
* Es decir, cada variable se eleva a todas las potencias entre 0 y p, y las combinaciones entre ellas no incluyen lostérminos elevados a p

```python
poly = PolynomialFeatures(6)
XX = poly.fit_transform(data2[:,0:-1])
print(XX.shape)
```

#### Función de coste regularizada
* Observar que se introduce el sumatorio de los parámetros $\theta$ ponderado por el valor del parámetro $\lambda$
    * Si $\lambda=0$ obtenemos la función de coste implementada en la primera parte de la práctica
    * Si $\lambda>0$ la función de coste trata de disminuir el valor de los parámetros $\theta$ del modelo
        * A mayor valor de $\lambda$ menores valores de los parámetros $\theta$ obtenedremos ya que trataremos de minimizarlos
    
#### $$ J(\theta) = \frac{1}{m}\sum_{i=1}^{m}\big[-y^{(i)}\, log\,( h_\theta\,(x^{(i)}))-(1-y^{(i)})\,log\,(1-h_\theta(x^{(i)}))\big] + \frac{\lambda}{2m}\sum_{j=1}^{n}\theta_{j}^{2}$$
#### Función de coste vectorizada
#### $$ J(\theta) = \frac{1}{m}\big((\,log\,(g(X\theta))^Ty+(\,log\,(1-g(X\theta))^T(1-y)\big) + \frac{\lambda}{2m}\sum_{j=1}^{n}\theta_{j}^{2}$$

Definir la función de coste que inluye la regularización. Para ello se debe seguir usando la función sigmoide definida en la primera parte de la práctica.

```python
def funcionCosteRegularizada(theta, valorLambda, XX, y):
    m = y.size
    h = sigmoide(XX.dot(theta))
    
    J = -1.0*(1.0/m)*(np.log(h).T.dot(y)+np.log(1-h).T.dot(1-y)) + (valorLambda/(2.0*m))*np.sum(np.square(theta[1:]))
    
    if np.isnan(J[0]):
        return(np.inf)
    return(J[0])
```

Recordamos las derivadas parciales de la regresión logística regularizada

#### Derivada parcial

#### $$ \frac{\partial J(\theta)}{\partial\theta_{j}} = \frac{1}{m}\sum_{i=1}^{m} ( h_\theta (x^{(i)})-y^{(i)})x^{(i)}_{j} + \frac{\lambda}{m}\theta_{j}$$ 
####  Derivada parcial vectorizada
#### $$ \frac{\partial J(\theta)}{\partial\theta_{j}} = \frac{1}{m} X^T(g(X\theta)-y) + \frac{\lambda}{m}\theta_{j}$$
Nota: el valor del término independiente $\theta_{0}$ no debe ser normalizado.

Definir la función que realiza la derivada parcial de los parámetros de la regresión logística regularizada en base al error cometido. Para ello se debe aplicar la ecuación de la derivada parcial vectorizada.

```python
def derivadaParcialRegularizada(theta, valorLambda, XX, y):
    m = y.size
    h = sigmoide(XX.dot(theta.reshape(-1,1)))
      
    derParcial = (1.0/m)*XX.T.dot(h-y) + (valorLambda/m)*np.r_[[[0]],theta[1:].reshape(-1,1)]
        
    return(derParcial.flatten())
```

Ejercicio:
* Inicializa los valores iniciales de la regresión logística a una columna de ceros
* Llama a la función que calcula el coste del modelo con los datos transformados (XX) mediante los polinomios aprendidos previamente y el parámetro $\lambda=1.0$.

<!-- grade_id: cell-51a12fa38d5c491a, points: 0, grade: false, solution: true -->
```python
# Columna de ceros como valores inciales de los parametros del modelo
# thetaInicial = #<RELLENAR>


# Llamada a la funcion que calcula el coste del modelo inicial
# coste = #<RELLENAR>

print(coste)
```

<!-- grade_id: cell-b04be24d707c9d2c, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(coste, 5), 0.69315)
```

Por último vamos a analizar el efecto del parámetro $\lambda$ en los resultados obtenidos. Para ello vamos a analizar 3 valores: 0, 1 y 100. Para cada valor hay que realizar el siguiente proceso:
* Llamar a la función minimize (descenso por gradiente) pasando como argumentos de entrada:
    * La función que calcula el coste del modelo
    * Los valores iniciales del modelo que acabamos de asignar
    * Una tupla que contiene: el valor de $\lambda$, los datos transformados con los polinomios (XX) y las clases de los 118 ejemplos (parámetro args)
    * La función que realiza la derivada parcial de los parámetros del modelo (parámetro jac)
    * El número máximo de iteraciones que lo fijamos a 3000
* Realizar la predicción de los ejemplos del histórico (datos guardados en XX) con el modelo aprendido (res2.x) (llamada a la función prediccion definida anteriormente)
* Calcular el porcentaje de acierto del modelo (se calcula como se ha realizado en la primera parte de la práctica)
* Mostrar la gráfica con la frontera de decisión generada por el modelo aprendido (llamada a la función que muestra los datos y sobre la cual mostramos la frontera de decisión generada)
    * Como caption mostramos el valor de $\lambda$, el porcentaje de acierto
* Mostrar los valores de los parámetros del modelo

Podéis observar los siguientes hechos:
* Con $\lambda=0$ el modelo obtenido es muy complejo puesto que está muy ajustado a los datos con los que ha sido aprendido. El porcentaje de acierto es muy alto (91.53%) y los valores de los parámetros $\theta$ son muy altos
* Con $\lambda=1$ el modelo obtenido es más sencillo y ofrece un buen balance entre complejidad y porcentaje de acierto (83.05%). Los valores de los parámetros $\theta$ son más bajos que antes.
* Con $\lambda=100$ el modelo obtenido es muy sencillo pero no se ajusta bien a los datos de entrenamiento y por tanto se obtiene un porcentaje de acierto muy bajo para este problema (61.02%). Los valores de los parámetros $\theta$ son muy bajos.

<!-- grade_id: cell-7fac546456c2f4be, points: 0, grade: false, solution: true -->
```python
fig, axes = plt.subplots(1,3, sharey = True, figsize=(17,5))

# Fronteras de decisión
# Lambda = 0 : No se realiza regularización --> modelo muy flexible: sobre-entrenado
# Lambda = 1 : Parece que es un buen modelo: simple y preciso
# Lambda = 100 : Demasiada regularización: mucho bías

for i, valorLambda in enumerate([0.0, 1.0, 100.0]):
    # Optimización de los parámetros del modelo en base al coste
    res2 = minimize(funcionCosteRegularizada, thetaInicial, args=(valorLambda, XX, y), jac=derivadaParcialRegularizada, options={'maxiter':3000})
    
    # Obtener las predicciones del modelo (res2.x) para los datos transformados con los polinomios (XX)
#     predicciones = #<RELLENAR>

    # Calcular el porcentaje de acierto como hemos realizado en la primera parte de la práctica
#     accuracy = #<RELLENAR>

    # Gráfica que muestra los ejemplos
    muestraGrafica(data2, 'Test 1 del microchip', 'Test 2 del microchip', 'Aceptado', 'Rechazado', axes.flatten()[i])
    
    # Mostramos la frontera de decisión
    x1_min, x1_max = X[:,0].min(), X[:,0].max(),
    x2_min, x2_max = X[:,1].min(), X[:,1].max(),
    print(x1_min, x1_max)
    print(x2_min, x2_max)
    xx1, xx2 = np.meshgrid(np.linspace(x1_min, x1_max), np.linspace(x2_min, x2_max))
#     print np.c_[xx1.ravel(), xx2.ravel()][:3]
#     print poly.fit_transform(np.c_[xx1.ravel(), xx2.ravel()]).dot(res2.x)[:5]
    h = sigmoide(poly.fit_transform(np.c_[xx1.ravel(), xx2.ravel()]).dot(res2.x))
    h = h.reshape(xx1.shape)
#     print (h>0.49).sum()
        
    axes.flatten()[i].contour(xx1, xx2, h, [0.5], linewidths=1, colors='g');       
    axes.flatten()[i].set_title('Porcentaje de acierto {}% con Lambda = {}'.format(np.round(accuracy, decimals=2), valorLambda))
    
    print('Valores de theta para lambda {}: {}'.format(valorLambda, res2.x))
```

```python

```