## Ejercicio 1 - Regresión lineal

- [Regresión lineal con una variable](#Regresión-lineal-con-una-variable)
- [Descenso por gradiente](#Descenso-por-gradiente)

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
from mpl_toolkits.mplot3d import Axes3D
from nose.tools import assert_equal

pd.set_option('display.notebook_repr_html', False)
pd.set_option('display.max_columns', None)
pd.set_option('display.max_rows', 150)
pd.set_option('display.max_seq_items', None)
 
#%config InlineBackend.figure_formats = {'pdf',}
%matplotlib inline  
```

### Regresión lineal con una variable

En primer lugar se debe leer los datos almacenados en el fichero `ex1data.txt` mediante la **función loadtxt de NumPy**. Los datos contenidos en este fichero corresponden a un problema en el que se debe predecir el beneficio de un camión de comida (food truck) en base al número de habitantes de una ciudad. En este fichero está almacenado el histórico de los camiones de comida de una franquicia. La primera columna corresponde al número de habitantes de una ciudad (en decenas de miles) y la segunda columna corresponde a los beneficios obtenidos con el camión de comida (en decenas de miles).

Mostrar los 5 primeros registros para comprobar que la lectura se haya realizado correctamente.

<!-- grade_id: cell-870c2e63684d2752, points: 0, grade: false, solution: true -->
```python
data = np.loadtxt('ex1data1.txt', delimiter=',')
# YOUR CODE HERE
raise NotImplementedError()
```

A continuación debéis crear los datos de entrada (X) y de salida (y) para poder realizar el aprendizaje de los parámetros de la regresión lineal ($\theta$). Recordar que la primera columna de los datos de entrada debe estar compuesta por unos para multiplicar al término independiente de la regresión lineal. Por tanto, la variable X estará compuesta por dos columnas: una de unos y otra con los datos correspondientes al número de habitantes de las ciudades. La variable y estará compuesta por los beneficios de los camiones de comida en las diferentes ciudades.

<!-- grade_id: cell-4399eaac487c2c0f, points: 0, grade: false, solution: true -->
```python
# La función hstack de NumPy concatena dos arrays horizontalmente
    # Debéis crear una columna de unos y concaternala con la primera columna de data
    # Debéis cambiar las dimensiones de ambos arrays para que sean una columna (dimensiones (97L, 1L))
# La función reshape de NumPy cambia las dimensiones de un array, para crear una sola columna se debe utilizar reshape(-1,1)
    # El -1 calcula el número de filas adecuado para el número de columnas especificado (en este caso 1)

# X = 
# YOUR CODE HERE
raise NotImplementedError()

# y = 
# YOUR CODE HERE
raise NotImplementedError()
```

Vamos a mostrar una gráfica de puntos para visualizar la distribución de las instancias del problema. De esta forma podemos comprobar si es posible que el modelo de regresión lineal se ajuste bien a los datos o no.

```python
plt.scatter(X[:,1], y, s=30, c='r', marker='x', linewidths=1)
plt.xlim(4,24)
plt.xlabel('Numero de habitantes de las ciudades en decenas de miles.')
plt.ylabel('Beneficio en decenas de miles de dolares.');
plt.show()
```

#### Descenso por gradiente

A continuación se debe programar la función de coste para poder aplicar el algoritmo de aprendizaje y que los parámetros de la regresión lineal se ajusten a los datos de tal forma que se minimice dicho coste.

Recordar que la función de coste es: $J(\theta) = \frac{1}{2m}\sum_{i=1}^m \left(h_{\theta}\left(x^{(i)}\right)-y^{(i)}\right)^2$

La instrucción X.dot(theta) realiza el producto matricial de los datos de entrada por los parámetros de la regresión. Su resultado es la predicción del sistema para los datos de entrada: $h_{\theta}\left(x^{(i)}\right)$

<!-- grade_id: cell-fb37a90b945f176b, points: 0, grade: false, solution: true -->
```python
def computeCost(X, y, theta):
    m = y.size
    J = 0
    
    h = X.dot(theta)
    
#     J = 
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return(J)
```

Vamos a comprobar si la función de coste está bien hecha. Para ello debéis crear los valores iniciales de los parámetros de la regresión lineal como un array columna de ceros (lista de listas de un elemento). Luego, debéis llamar a la función computeCost con los datos de entrada (X), de salida (y) y los parámetros del modelo ($\theta$). Almacenar el resultado en la variable resultado.

Imprimir en pantalla el resultado.

<!-- grade_id: cell-e9d1853ce8d3f777, points: 0, grade: false, solution: true -->
```python
# theta = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()

# resultado = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()
print(resultado)
```

<!-- grade_id: cell-112fd6c583614235, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(resultado, 5), 32.07273)
```

A continuación se muestra la implementación de la función que implementa el descenso por gradiente para realizar el aprendizaje de los parámetros del modelo. Esta función recibe como parámetros de entrada los valores de entrada, de salida, los valores iniciales del modelo, el valor que establece la velocidad de convergencia del algoritmo ($\alpha$) y el número de iteraciones a realizar por el algoritmo.

Como salida devuelve los parámetros aprendidos (variable theta) así como el historial de la evolución de error cometido con cada configuración de valores del modelo (variable J_history). Fijaros que esta última variable almacena el resultado de la función de coste implementada anteriormente para los valores del modelo que han sido actualizados en la instrucción anterior. 

```python
def gradientDescent(X, y, theta, alpha=0.01, num_iters=1500):
    m = y.size
    J_history = np.zeros(num_iters)
    
    for iter in np.arange(num_iters):
        h = X.dot(theta)
        theta = theta - alpha*(1.0/m)*(X.T.dot(h-y))
        J_history[iter] = computeCost(X, y, theta)
    return(theta, J_history)
```

Vamos a realizar el aprendizaje. Para ello se debe llamar a la función gradientDescent con los valores de entrada (X), los de salida (y) y los valores iniciales de los parámetros del modelo ($\theta$), inicializados a cero anteriormente.

A continuación se muestran por pantalla y se comprueba si son correctos.

Finalmente, se muestra la gráfica de la evolución del error con respecto a las iteraciones del descenso por gradiente.

<!-- grade_id: cell-8fbd74952d022ef1, points: 0, grade: false, solution: true -->
```python
# theta para minimizar la función de coste J
# theta , Cost_J = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()
print('theta: ',theta.ravel())

plt.plot(Cost_J)
plt.ylabel('Cost J')
plt.xlabel('Iterations');
```

<!-- grade_id: cell-4500b01f95d6e5f2, points: 6, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: round(ind, 5), list(theta.ravel()))), [-3.63029,  1.16636])
```

Ahora vamos a mostrar la recta aprendida. 

Para ello debéis crear una array de valores comprendidos entre 5 y 23 con incrementos de 1 en 1 y almacenarlo en la variable xx. 

A continuación evaluar todos estos valores utilizando los valores del modelo aprendidos anteriormente (variable theta) y almacenarlos en la variable yy. Se debe aplicar la ecuación de la recta: $y = n+m*x$.

Por último se muestra la recta aprendida en la figura que también incluye los puntos a partir de los que se ha realizado el aprendizaje.


<!-- grade_id: cell-4ae41c92bd6eeab2, points: 0, grade: false, solution: true -->
```python
# xx = #<RELLENAR>
# YY = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()

# Mostramos el descenso por gradiente
plt.scatter(X[:,1], y, s=30, c='r', marker='x', linewidths=1)
plt.plot(xx,yy, label='Regresion lineal (Descenso por gradiente)')

plt.xlim(4,24)
plt.xlabel('Numero de habitantes de las ciudades en decenas de miles.')
plt.ylabel('Beneficio en decenas de miles de dolares.');
plt.legend(loc=4);
```

Una vez realizado el proceso de aprendizaje con funciones definidas por nosotros vamos a pasar a usar la librería ofrecida por Python para realizar tareas de aprendizaje automático. Esta librería se llama scikit-learn y se puede visitar en la siguiente URL: http://scikit-learn.org/stable/. Iremos aprendiendo a usarla a lo largo de la asignatura. 

Hoy vamos a comenzar utilizando las funcionalidades relacionadas con la regresión lineal. Podéis consular toda la información en la siguiente URL: http://scikit-learn.org/stable/modules/linear_model.html

* Para ello hay que importar la clase LinearRegression del paquete linear_model ofrecido en scikit-learn (la librería se llama sklearn). 
* Una vez importada la clase debemos llamar al constructor para inicializar el modelo de regresión lineal. 
* Después hay que realizar el aprendizaje y para ello se utiliza la función fit, a la que se le pasan los datos de entrada (X) y las salidas correspondientes (y). Esta función automáticamente realiza el aprendizaje y modifica internamente los parámetros del modelo (intercept_ y coef_ para el término independiente y los $\theta$ asociados a cada variable del problema).
* Para clasificar nuevos datos tenemos dos opciones
    * Utilizar intercept_ y coef_ aplicando la ecuación de la regresión lineal
    * Utilizar la función predict, a la que se le pasa los datos de entrada de las instancias a predecir y devuelve las predicciones.

En el código mostrado a continuación se ve un ejemplo del uso de esta librería con comentarios sobre qué se realiza en cada instrucción de código. Finalmente se muestra en una gráfica tanto la recta aprendida anteriormente como la recta aprendida con la librería scikit-learn.

```python
# Se importa de la libreria scikit-learn de python (sklearn) el paquete correspondiente a la regresión lineal
from sklearn.linear_model import LinearRegression

# Se incializa el modelo llamando al constructor de la regresión lineal: todos los parámetros se asignan a sus valores por defecto
regr = LinearRegression()
# Se entrena el modelo (aprendizaje) utilizando la variable de entrada sin los unos X[:,1] y la variable de salida (y)
# La función ravel devuelve una lista (si la entrada es una lista de listas coge todos los elementos y los almacena en la lista devuelta)
# En la función reshape el valor -1 significa que el tamaño de esa dimensión se calcula automáticamente en base a la otra
    # En este caso hacemos una columna de tantas filas como instancias tengan los datos
regr.fit(X[:,1].reshape(-1,1), y.ravel())
# Se evalúan los valores creados anteriormente para comparar ambos modelos aprendidos
prediccionesOpcion1 = regr.intercept_+regr.coef_*xx
prediccionesOpcion2 = regr.predict(xx.reshape(-1,1))
plt.plot(xx, prediccionesOpcion2, label='Regresion lineal (Scikit-learn)')
# Se muestra los datos a partir de los cuales se ha aprendido
plt.scatter(X[:,1], y, s=30, c='r', marker='x', linewidths=1)
# Se muestra la recta aprendida anteriormente para comparar
plt.plot(xx,yy, label='Regresion lineal (Descenso por gradiente)')

plt.xlim(4,24)
plt.xlabel('Numero de habitantes de las ciudades en decenas de miles.')
plt.ylabel('Beneficio en decenas de miles de dolares.');
plt.legend(loc=4);
```

Una vez aprendido el modelo de regresión lineal podemos utilizarlo para predecir datos que no han sido utilizados durante el entrenamiento. En este caso debéis predecir el valor de dos ciudades: la primera tiene 35.000 habitantes y la segunda 70.000 habitantes. Recordar que tanto el número de habitantes como el beneficio con el que se ha aprendido el modelo de regresión lineal está expresado en decenas de miles de habitantes.

Predecir los beneficios de ambas ciudades utilizando los valores de los parámetros aprendidos por las funciones definidas por nosotros.


<!-- grade_id: cell-9edd08902ad7f747, points: 0, grade: false, solution: true -->
```python
# Predecir el beneficio para dos ciudades de 35000 y 70000 habitantes, respectivamente
# ciudad1 = #<RELLENAR>
# ciudad2 = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()
print(ciudad1)
print(ciudad2)
```

<!-- grade_id: cell-830cd9802ce15742, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(ciudad1[0], 5), 4519.76787)
assert_equal(round(ciudad2[0], 5), 45342.45013)
```

Finalmente, os dejamos un código que se utiliza para visualizar el comportamiento de la función de coste ($J$) en función de los parámetros del modelo ($\theta_1$ y $\theta_2$). Se crear dos gráficas, una para mostrar el comportamiento de la función de error en 3 dimensiones y la otra es una gráfica de contornos en la que se observa el nivel de error en cada punto.

```python
# Se crean 50 puntos en cada dimensión: en el rango [-10, 10] para la población y [-1, 4] para los beneficios
B0 = np.linspace(-10, 10, 50)
B1 = np.linspace(-1, 4, 50)
# Se crean todas las combinaciones de valores de los puntos creados anteriormente
xx, yy = np.meshgrid(B0, B1, indexing='xy')
# Se crea una matriz donde vamos a almacenar el coste de la predicción (error) para cada combinación
Z = np.zeros((B0.size,B1.size))

# Se calcula el error de cada combinación
for (i,j),v in np.ndenumerate(Z):
    Z[i,j] = computeCost(X,y, theta=[[xx[i,j]], [yy[i,j]]])

# Se crea una figura con dos plots: uno para cada tipo de gráfico
fig = plt.figure(figsize=(15,6))
ax1 = fig.add_subplot(121)
ax2 = fig.add_subplot(122, projection='3d')

# Gráfica de la izquierda: contornos (nivel de error)
CS = ax1.contour(xx, yy, Z, np.logspace(-2, 3, 20), cmap=plt.cm.jet)
ax1.scatter(theta[0],theta[1], c='r')

# Gráfica de la izquierda: comportamiento de la función de coste
ax2.plot_surface(xx, yy, Z, rstride=1, cstride=1, alpha=0.6, cmap=plt.cm.jet)
ax2.set_zlabel('Coste')
ax2.set_zlim(Z.min(),Z.max())
ax2.view_init(elev=15, azim=230)

# Se asignan las etiquetas de los ejes de las dos gráficas
for ax in fig.axes:
    ax.set_xlabel(r'$\theta_0$', fontsize=17)
    ax.set_ylabel(r'$\theta_1$', fontsize=17)
```

```python

```