## Ejercicio 2 - Regresión lineal

- [Regresión lineal con varias variables](#Regresión-lineal-con-varias-variables)
- [Descenso por gradiente](#Descenso-por-gradiente)

```python
# %load ../../standard_import.txt
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
from nose.tools import assert_equal

pd.set_option('display.notebook_repr_html', False)
pd.set_option('display.max_columns', None)
pd.set_option('display.max_rows', 150)
pd.set_option('display.max_seq_items', None)
 
#%config InlineBackend.figure_formats = {'pdf',}
%matplotlib inline  
```

### Regresión lineal con varias variables

En el segundo ejercicio vamos a tratar un problema de regresión lineal de varias variables de entrada. En este caso van a ser 2 pero podría usarse para cualquier número de variables de entrada.

Para ello, en primer lugar se debe leer los datos almacenados en el fichero ex1data2.txt mediante la función loadtxt de NumPy. Los datos contenidos en este fichero corresponden a un problema en el que se debe predecir el valor de una casa en base a su pies cuadrados y su número de habitaciones. En este fichero está almacenado el histórico de las casas valoradas en la actualidad. La primera columna corresponde al número de pies cuadrados, la segunda al número de habitaciones y la tercera a la valoración realizada.


```python
data = np.loadtxt('ex1data2.txt', delimiter=',')
```

En este caso debemos estandarizar las variables ya que los rangos de las mismas son muy diferentes e introduciríamos un bias en favor del número de pies cuadrados.

Para estandarizar los datos tenéis que definir una función llamada estandarizar que para cada variable debéis hacer el siguiente proceso:
* Calcular la media de los valores de la variable para todas las instancias
* Calcular la desviación estándar de los valores de la variable para todas las instancias
* A cada valor se le debe restar la media y el resultado dividirlo por la desviación estándar
* Devolver los datos estandarizados así como la media y la desviación estándar calculadas

<!-- grade_id: cell-6be6ad7333dabbd4, points: 0, grade: false, solution: true -->
```python
def estandarizar(datos):
    # Debemos crear una copia de los datos para que no se pierdan sus valores originales
    datosEstandarizados = np.copy(datos)
    # Calculamos la media y desviación por columnas y estandarizamos los datos
    # YOUR CODE HERE
    raise NotImplementedError()
    return datosEstandarizados, media, desv
```

Para estandarizar nuevos datos de test se debe realizar el mismo proceso que en el caso anterior pero, como son datos nuevos (datos de test), tenemos que utilizar la media y la desviación estándar obtenidas con los datos de entrenamiento. De esta forma la estandarización se realizará para todos los datos en las mismas condiciones. 

Definir una función que estandarice nuevos datos de test. Para ello se aplica la misma ecuación que para la función anterior pero utilizando la media y la desviación estándar pasadas como parámetros de entrada (serán las obtenidas con el conjunto de entrenamiento). La función devuelve los nuevos datos estandarizados.

<!-- grade_id: cell-52ec3992ccb6347f, points: 0, grade: false, solution: true -->
```python
def estandarizarTest(datos, media, desv):
    # Debemos crear una copia de los datos para que no se pierdan sus valores originales
    datosEstandarizados = np.copy(datos)
    # Estandarizamos los datos
    # YOUR CODE HERE
    raise NotImplementedError()
    return datosEstandarizados
```

Ahora debéis utilizar la función estandarizar para realizar la estandarización de los datos de entrada almacenados en la variable data y almacenar los datos estandarizados en una variable llamada entrada así como la media y la desviación estándar calculadas con estos datos de entrenamiento.

<!-- grade_id: cell-753d91c4814fb4e2, points: 0, grade: false, solution: true -->
```python
# entrada, media, desv = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()
```

A continuación debéis crear los datos de entrada (X) y de salida (y) para poder realizar el aprendizaje de los parámetros de la regresión lineal ($\theta$). Recordar que la primera columna de los datos de entrada debe estar compuesta por unos para multiplicar al término independiente de la regresión lineal. Por tanto, la variable X estará compuesta por tantas columnas como variables de entrada más una (la de unos). En este caso serán 3 columnas.

<!-- grade_id: cell-cee384353caa5374, points: 0, grade: false, solution: true -->
```python
# # X = #<RELLENAR>
# y = #<RELLENAR> 
# YOUR CODE HERE
raise NotImplementedError()
print(X[:5,:])
print(y[:5])
```

#### Descenso por gradiente

A continuación debéis pensar si es necesario cambiar la función de coste programada en el primer ejercicio.

Recordar que la función de coste es: $J(\theta) = \frac{1}{2m}\sum_{i=1}^m \left(h_{\theta}\left(x^{(i)}\right)-y^{(i)}\right)^2$


<!-- grade_id: cell-43a2f202b69805e6, points: 0, grade: false, solution: true -->
```python
def computeCost(X, y, theta):
    m = y.size
    J = 0
    
    h = X.dot(theta)
    
#     J = #<RELLENAR>
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return(J)
```

Vamos a comprobar si la función de coste está bien hecha. 

Recordar que debéis crear los valores iniciales de los parámetros de la regresión lineal como un array columna de ceros (lista de listas de un elemento). Luego, debéis llamar a la función computeCost con los datos de entrada (X), de salida (y) y los parámetros del modelo ($\theta$). Almacenar el resultado en la variable resultado.

Imprimir en pantalla el resultado y comprobar si es correcto.

<!-- grade_id: cell-8cd2de40948e63bc, points: 0, grade: false, solution: true -->
```python
# theta = #<RELLENAR>
# resultado = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()
print(resultado)
```

<!-- grade_id: cell-c7b67ed5de618d18, points: 3, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(resultado, 1), 65591548106.5)
```

Ahora debéis pensar si es necesario modificar la función que realiza el descenso por gradiente para realizar el aprendizaje de los parámetros de la regresión lineal en base a los datos.

```python
def gradientDescent(X, y, theta, alpha=0.01, num_iters=1500):
    m = y.size
    J_history = np.zeros(num_iters)
    
    for iter in np.arange(num_iters):
        h = X.dot(theta)
        theta = theta - alpha*(1.0/m)*(X.T.dot(h-y))
        J_history[iter] = computeCost(X, y, theta)
    return(theta, J_history)
```

Realizar el aprendizaje llamando a la función gradientDescent con los valores de entrada (X), los de salida (y) y los valores iniciales de los parámetros del modelo ($\theta$), inicializados a cero anteriormente.

A continuación se muestran por pantalla y se comprueba si son correctos.

Finalmente, se muestra la gráfica de la evolución del error con respecto a las iteraciones del descenso por gradiente.

<!-- grade_id: cell-d931b676f264556d, points: 0, grade: false, solution: true -->
```python
# theta for minimized cost J
# theta, Cost_J = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()
print('theta: ',theta.ravel())

plt.plot(Cost_J)
plt.ylabel('Cost J')
plt.xlabel('Iterations');
```

<!-- grade_id: cell-92efc96ef80b48a3, points: 6, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(list(map(lambda ind: round(ind, 5), list(theta.ravel()))), [340412.56301,  109370.05670,   -6500.61510])
```

Realizar el aprendizaje utilizando el paquete de regresión lineal contenido en la librería scikit-learn de Python.

<!-- grade_id: cell-1746eb9c846022c3, points: 0, grade: false, solution: true -->
```python
#se importa de la libreria scikit-learn de python (sklearn) el paquete correspondiente a la regresión lineal
from sklearn.linear_model import LinearRegression

# Se incializa el modelo llamando al constructor de la regresión lineal: todos los parámetros se asignan a sus valores por defecto
regr = LinearRegression()
# Se entrena el modelo (aprendizaje) utilizando la variable de entrada sin los unos X[:,1] y la variable de salida (y)
# regr.fit(#<RELLENAR>)
# YOUR CODE HERE
raise NotImplementedError()
```

Predecir el valor de una casa de 1500 pies cuadrados cuadrados y 4 habitaciones. Como es un nuevo dato (test) se debe estandarizar con la función de estandarización de nuevos datos y utilizando la media y la desviación estándar de los datos de entrenamiento calculados previamente.

Predecir el valor de la casa utilizando el modelo aprendido con scikit-learn (función predict).


<!-- grade_id: cell-237c00519e6f591c, points: 0, grade: false, solution: true -->
```python
# Obtén la predicción para una casa de 1500 pies cuadrados y 4 habitaciones
# valorCasa = #<RELLENAR>
# YOUR CODE HERE
raise NotImplementedError()

print(valorCasa)
```

<!-- grade_id: cell-e40f1d3184f44787, points: 4, grade: true, solution: false -->
```python
# ESTA CELDA DARÁ ERROR SI EL RESULTADO NO ES CORRECTO
    # EN CASO CONTRARIO NO TENDRÁ SALIDA
assert_equal(round(valorCasa[0], 4), 263461.8441)
```

```python

```